## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc, to develop a new version of the function with an improved interface (the interface is the arguments that are passed to the function) with fewer arguments. Add appropriate checks inside your function on the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
def area(A, B, C): # A, B, C are the vertices of the triangle in the form (x, y)
    # assert that A, B and C are tuples
    try:
        assert type(A) == tuple
        assert type(B) == tuple
        assert type(C) == tuple
        assert len(A) == 2
        assert len(B) == 2
        assert len(C) == 2
    except AssertionError:
        raise ValueError('A, B and C must be tuples of length 2!')

    # assert that the coordinates are floating point numbers
    try:
        assert type(A[0]) == float
        assert type(A[1]) == float
        assert type(B[0]) == float
        assert type(B[1]) == float
        assert type(C[0]) == float
        assert type(C[1]) == float
    except AssertionError:
        raise ValueError('The coordinates must be floating point numbers!')

    numerator = (A[0] * (B[1] - C[1])) + (B[0] * (C[1] - A[1])) + (C[0] * (A[1] - B[1]))
    return abs(numerator / 2)

print(area((0.0, -2.0), (0.0, 2.0), (2.5, 0.0)))
```

Output:
```
5.0
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that, given only the vertex coordinates, computes the area of a simple polygon. Test your function for some simple shapes. Pay careful attention to the range of any loops. **Test for polygons without a vertex at the origin.**

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
import math

def calculate_polygon_area(vertices):
    n = len(vertices)
    extended_vertices = (*vertices, vertices[0]) # append the first vertex to the end of the tuple
    
    current_sum = 0.0 # current value of the sum as we loop through the tuple

    for i in range(n):
        current_sum += (extended_vertices[i][0] * extended_vertices[i + 1][1]) - (extended_vertices[i + 1][0] * extended_vertices[i][1])
    
    return abs(current_sum / 2)

# test with a quadrilateral
quadrilateral_vertices = ((2.5, 0.0), (1.0, 1.0), (-1.0, -1.0), (0.0, -2.0))
quadrilateral_area = calculate_polygon_area(quadrilateral_vertices)

assert math.isclose(quadrilateral_area, 4.75) # check against expected value
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match. Print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(v_1, v_2):
    """Return sum of two vectors"""
    # check that the vectors have the same length
    if len(v_1) != len(v_2):
        raise ValueError("The two vectors being added must be of equal length!")
    
    n = len(v_1)
    v_sum = [0.0] * n

    for i in range(n):
        v_sum[i] = v_1[i] + v_2[i]

    return v_sum
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
# tests ##
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Optional (advanced) 

Try writing a one-line version that performs the `sum_vector` operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
...
```

Output:
```
Ellipsis
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A Python dictionary from college abbreviation to name; and
1. A Python list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
college_name_to_abbreviation = {
    "Christ's": "CHR",
    "Churchill": "CHU",
    "Clare": "CL",
    "Clare Hall": "CLH",
    "Corpus Christi": "CC"
}

print("names to abbreviations:", college_name_to_abbreviation)

# invert the key: value relationship
college_abbreviation_to_name = {abbreviation: name for name, abbreviation in college_name_to_abbreviation.items()}

print("abbreviations to names:", college_abbreviation_to_name)

abbreviation_list = sorted(college_name_to_abbreviation.values())

print("sorted abbreviations:", *abbreviation_list)
```

Output:
```
names to abbreviations: {"Christ's": 'CHR', 'Churchill': 'CHU', 'Clare': 'CL', 'Clare Hall': 'CLH', 'Corpus Christi': 'CC'}
abbreviations to names: {'CHR': "Christ's", 'CHU': 'Churchill', 'CL': 'Clare', 'CLH': 'Clare Hall', 'CC': 'Corpus Christi'}
sorted abbreviations: CC CHR CHU CL CLH
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
...
```

Output:
```
Ellipsis
```