## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n): 
    "Compute the nth Fibonacci number using recursion"
    if n == 0:
        return 0
    elif n == 1:
        return 1
    elif n > 1:
        return f(n - 1) + f(n - 2)
    else:
        raise ValueError("n must be greater than or equal to 0!")
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
## tests ##

# Perform some tests    
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    """Evaluate polynomial function"""
    return x**5 / 10 + x**3 - 10 * x**2 + 4 * x + 7

# Create the function that performs the bisection:
def compute_root(f, x_0, x_1, tol, max_it):
    """Compute roots of a function using bisection"""

    # input data checks
    try:
        x_0 = float(x_0)
        x_1 = float(x_1)
        tol = float(tol)
        max_it = int(max_it)
    except ValueError:
        raise ValueError("x_0 must be a float, x_1 must be a float, tol must be a float & max_it must be an int")

    it = 0
    f_mid = tol + 1.0
    while abs(f_mid) > tol:
        if it >= max_it:
            raise RuntimeError("Max number of iterations exceeded!")
        # calculate midpoint
        x_mid = (x_0 + x_1) / 2

        # Evaluate function at (i) left end-point and at (ii) midpoint
        f_0 = f(x_0)
        f_mid = f(x_mid)
        
        # calculate the product of f_0 and f_mid
        product = f_0 * f_mid

        # check the sign of the product
        if product < 0.0:
            x_1 = x_mid # f changes sign between x_0 and x_mid
        else:
            x_0 = x_mid # f changes sign between x_mid and x_1

        # increment iteration counter
        it += 1

     
    return x_mid, f_mid, it
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
## tests ##

# Test with max_it = 30
x, f, num_it = compute_root(my_f, 0, 2, 1.0e-6, 30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, 0, 2, 1.0e-6, 20)
```